# CPSC 4300/6300-001 Applied Data Science (Fall 2020)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Dane Acena"
COLLABORATORS = ""
```

# CPSC4300/6300-001 Problem Set #1

In this assignment, you will create a NCAA College Football team dataset by scraping data from the Internet and perform some basic data manipulations. The assignment has four parts: the first three are auto-graded. Th elast part is manually graded by the TA. 

Before submission, make sure you run all the cells. Because the TA could not run your notebook, your grade will be only based on the output and the writing part of your submitted notebook. 

```python

```

## Part B: Scape the Rosters of Any NCAA Collge Football Team (50 points)

During the class, we have shown how to extract the roster of the Clemson Tiger's defense team, now we want to extend the work to apply to the entire team roster of any NCAA coolege football teams.

### Question 1. Wriet a function to computer the Team URL

Look at the URLs of server college football team at the ESPN site, generalize the pattern, and write a python function  `get_team_url(team_id, team_name)` to compute a url from team_id and team_name. (4 points)

<!-- grade_id: cell-e972833e6447cb14, points: 0, grade: false, solution: true -->
```python
def get_team_url(team_id, team_name):
    """
    Get the team URL for a NCAA college football team
    
    Args:
        team_id: the id of a NACC football team assigned by ESPN
        team_name: the name of the football team
        
    Returns:
        team_url: a string representing the URL of the team
    """
    url_team_name = team_name.replace(" ","-").lower()
    team_url = 'https://www.espn.com/college-football/team/roster/_/id/%s/%s' % (team_id,url_team_name)
    return team_url
#     raise NotImplementedError()
```

<!-- grade_id: cell-ce7ed3ddbacff2e2, points: 2, grade: true, solution: false -->
```python
team_id = 228
team_name = "Clemson Tigers"
team_url = get_team_url(team_id, team_name)
assert team_url == 'https://www.espn.com/college-football/team/roster/_/id/228/clemson-tigers'
```

<!-- grade_id: cell-a5c7f35497643616, points: 2, grade: true, solution: false -->
```python
# This cell contains some hidden test cases.
```

### Question 2. Create a BeautifulSoup Object

Complete the following code to create a BeautifulSoup Object froma team url and name the object `soup`. (4 points)

<!-- grade_id: cell-affd9982ef36210b, points: 0, grade: false, solution: true -->
```python
url = get_team_url(228, "Clemson Tigers")

import requests
from bs4 import BeautifulSoup

req = requests.get(url)
soup = BeautifulSoup(req.content, 'html.parser')
# raise NotImplementedError()
type(soup)
```

Output:
```
bs4.BeautifulSoup
```

<!-- grade_id: cell-001d464a2c8f9785, points: 2, grade: true, solution: false -->
```python
assert (type(soup) == BeautifulSoup)
```

<!-- grade_id: cell-003e55323653f7bb, points: 2, grade: true, solution: false -->
```python
assert soup.find('title').text.startswith('Clemson')
```

### Question 3. Find the Team Name

Write a statement to find a tag whose text is formatted as "{Team Name} Rosters" on the page. The {Team Name} would be the actual team name like 'Clemson Tigers'. (2 points)

<!-- grade_id: cell-159f19bec0f548a6, points: 0, grade: false, solution: true -->
```python
team_tag = soup.find('h1', string='%s Roster' % team_name)
# raise NotImplementedError()
print(type(team_tag))
print(team_tag)
```

Output:
```
<class 'bs4.element.Tag'>
<h1 class="headline headline__h1 dib">Clemson Tigers Roster</h1>
```

<!-- grade_id: cell-0313f4fbe0319355, points: 2, grade: true, solution: false -->
```python
assert str(type(team_tag)) == "<class 'bs4.element.Tag'>" and team_tag.text.endswith("Roster")
```

### Question 5. Search for Team Groups

A football team roster is usually divied into three groups: *Offense*, *Defense*, and *Special Teams*. Write some statements to find the tags for all these three groups and save the tags into a variable called groups. (2 points)

<!-- grade_id: cell-7f4a507853f8b142, points: 0, grade: false, solution: true -->
```python
groups = soup.find_all(class_="Table__Title")
# raise NotImplementedError()
print(len(groups), "groups")
for g in groups:
    print(g)
```

Output:
```
3 groups
<div class="Table__Title">Offense</div>
<div class="Table__Title">Defense</div>
<div class="Table__Title">Special Teams</div>
```

<!-- grade_id: cell-42d807cdb2edf390, points: 1, grade: true, solution: false -->
```python
assert ','.join(sorted([g.text for g in groups])) == 'Defense,Offense,Special Teams'
```

<!-- grade_id: cell-71dc4780c401038e, points: 1, grade: true, solution: false -->
```python
assert len(groups) == 3
```

### Question 6. Find the Column Names

Complete the code segment to extract the names of the columns in a table rows and save those names into a list named col_names. (4 points)

<!-- grade_id: cell-2cbb21a90c5b9f64, points: 0, grade: false, solution: true -->
```python
t = soup.find(class_="Table__Title", string="Offense").find_next("table")
col_names = [column.get_text() for column in t.find_all('th')]
# table_header = soup.find_all(class_="Table__TH")
# # raise NotImplementedError()
# for columns in table_header:
#     col_names.append(columns.text)

# table_header = soup.find(class_="Table__Title", string="Offense").find_next("table")
# for columns in t.thead.tr.children:
#     col_names.append(columns.text)

col_names
```

Output:
```
['', 'Name', 'POS', 'HT', 'WT', 'Class', 'Birthplace']
```

<!-- grade_id: cell-ba933afaef916cfd, points: 2, grade: true, solution: false -->
```python
assert "Name" in col_names and "Class" in col_names
```

<!-- grade_id: cell-cf92003b92fab9c5, points: 2, grade: true, solution: false -->
```python
assert type(col_names) is list and len(col_names) >= 6
```

### Question 7. Find the Attributes of a Player

Write code segment to extract the attributes of the first defense player and save them into a dict. The player should have the attributes as shown in the following example.
```
player_name: Lannden Zanders
player_id: 4428971
player_no: 36
POS: S
HT: 6' 1"
WT: 195 lbs
Class: SO
Birthplace: Shelby, NC
```
(10 points)

<!-- grade_id: cell-667e97ed5447158a, points: 0, grade: false, solution: true -->
```python
import re

player = {}
players = []
player_attr = soup.find(class_="Table__Title", string="Defense").find_next('table')

def is_player_row(tag):
    if tag.name == 'tr' and tag.has_attr('data-idx'):
        return True
    return False
    
for c in player_attr.tbody.find_all(is_player_row):
    tds = c.find_all('td')
    urls = c.find_all('a', href=True)
    cols = []   
    for td in tds:
        cols.append(td.get_text())
    for url in urls:
        if (url['href']) not in cols:
            cols.append(url['href'])
    players.append(cols)

player_name = "".join(re.split("[^a-zA-Z ]*", players[0][1]))
player_no = "".join(re.split("[^0-9]*", players[0][1]))

player_url = players[0][7]
player_id = "".join(re.split("[^0-9]*", player_url))
player = {
    "player_name": player_name,
    "player_id": player_id,
    "player_no": player_no,
    "POS": players[0][2],
    "HT": players[0][3],
    "WT": players[0][4],
    "Class": players[0][5],
    "Birthplace": players[0][6]
}
print(player)
# raise NotImplementedError()
```

Output:
```
{'player_name': 'Bryan Bresee', 'player_id': '4428988', 'player_no': '11', 'POS': 'DL', 'HT': '6\' 5"', 'WT': '300 lbs', 'Class': 'FR', 'Birthplace': 'Damascus, MD'}
```

```python
for k, v in player.items():
    print("{}: {}".format(k, v))

```

Output:
```
player_name: Bryan Bresee
player_id: 4428988
player_no: 11
POS: DL
HT: 6' 5"
WT: 300 lbs
Class: FR
Birthplace: Damascus, MD
```

<!-- grade_id: cell-7c574fd6f9413477, points: 1, grade: true, solution: false -->
```python
assert type(player) is dict
```

<!-- grade_id: cell-75b251071e223227, points: 2, grade: true, solution: false -->
```python
assert len(player) == 8
```

<!-- grade_id: cell-9c5f3fdbcd88b09b, points: 5, grade: true, solution: false -->
```python
assert ','.join(sorted(player.keys())) == 'Birthplace,Class,HT,POS,WT,player_id,player_name,player_no'
```

<!-- grade_id: cell-7071bacb49ae8bef, points: 1, grade: true, solution: false -->
```python
assert int(player['player_no']) > 0 and int(player['player_no']) < 200
```

<!-- grade_id: cell-232a559e5c4360cc, points: 1, grade: true, solution: false -->
```python
assert int(player['WT'].split()[0]) > 150 and player['WT'].split()[1] == 'lbs'
```

### Question 8. Abstract the tasks into a function

Since you don't want to repeat the above steps for every player and every team, it is necessary to abstract the above scraping task into a function which you can reuse over and over. (10 points)

<!-- grade_id: cell-3de336d8b0a64126, points: 0, grade: false, solution: true -->
```python
import requests
from bs4 import BeautifulSoup

def is_player_row(tag):
    if tag.name == 'tr' and tag.has_attr('data-idx'):
        return True
    return False

def extract_team_roster(team_id, team_name):
    """
    Extract the team roster for a NCAA college football team
    
    Args:
        team_id: the id of a NACC football team assigned by ESPN
        team_name: the name of the football team
        
    Returns:
        a list of dict representing the roster of the entire team.
    """
    roster = []
    url_team_name = team_name.replace(" ","-").lower()
    team_url = 'https://www.espn.com/college-football/team/roster/_/id/%s/%s' % (team_id,url_team_name)
#     print(team_url)
    
    req = requests.get(team_url)
    soup = BeautifulSoup(req.content, 'html.parser')
    
    t = soup.find(class_="Table__Title").find_next("table")
    col_names = [column.get_text() for column in t.find_all('th')]
    
    player = {}
    players = []
    i = 0
    player_groups = soup.find_all(class_="Table__Title")
    
    for group_roster in player_groups:
        player_group_attr = soup.find(class_="Table__Title",string=group_roster.text).find_next('table')
        for c in player_group_attr.tbody.find_all(is_player_row):
            
            tds = c.find_all('td')
            urls = c.find_all('a', href=True)
            cols = []   
            for td in tds:
                cols.append(td.get_text())
            for url in urls:
                if (url['href']) not in cols:
                    cols.append(url['href'])
            players.append(cols)
            
            
            player_name = "".join(re.split("[^a-zA-Z ]*", players[i][1]))
            player_no = "".join(re.split("[^0-9]*", players[i][1]))

            player_url = players[i][7]
            player_id = "".join(re.split("[^0-9]*", player_url))
            player = {
                "player_name": player_name,
                "player_id": player_id,
                "player_no": player_no,
                "POS": players[i][2],
                "HT": players[i][3],
                "WT": players[i][4],
                "Class": players[i][5],
                "Birthplace": players[i][6]
            }
            roster.append(player)
            i = i + 1
    
#     raise NotImplementedError()
#     print(roster)
    return roster

team_id = 150
team_name = "Duke Blue Devils"
team_roster = extract_team_roster(team_id, team_name)
```

```python
# test the function
import pprint

team_id = 150
team_name = "Duke Blue Devils"
team_roster = extract_team_roster(team_id, team_name)

pp = pprint.PrettyPrinter(indent=4)
pp.pprint(team_roster[:3])
```

Output:
```
[   {   'Birthplace': 'Grayson, GA',
        'Class': 'JR',
        'HT': '6\' 3"',
        'POS': 'QB',
        'WT': '235 lbs',
        'player_id': '4239994',
        'player_name': 'Chase Brice',
        'player_no': '8'},
    {   'Birthplace': 'Langhorne, PA',
        'Class': 'SO',
        'HT': '6\' 3"',
        'POS': 'QB',
        'WT': '190 lbs',
        'player_id': '4568231',
        'player_name': 'Jack Colyar',
        'player_no': '13'},
    {   'Birthplace': 'Venice, CA',
        'Class': 'FR',
        'HT': '6\' 3"',
        'POS': 'QB',
        'WT': '180 lbs',
        'player_id': '4429578',
        'player_name': 'Luca Diamont',
        'player_no': '2'}]
```

<!-- grade_id: cell-895004dba91372be, points: 1, grade: true, solution: false -->
```python
type(team_roster) is list
```

Output:
```
True
```

<!-- grade_id: cell-018fc6188aad6ea2, points: 1, grade: true, solution: false -->
```python
assert len(team_roster) >= 80
```

<!-- grade_id: cell-a5e0daeb322a2166, points: 4, grade: true, solution: false -->
```python
assert ','.join(sorted(team_roster[50].keys())) == 'Birthplace,Class,HT,POS,WT,player_id,player_name,player_no'
```

<!-- grade_id: cell-fdf01ad39dd2e479, points: 2, grade: true, solution: false -->
```python
import numpy as np
classes = np.array([p['Class'] for p in team_roster])
class_labels, class_counts = np.unique(classes, return_counts=True)
assert ''.join(class_labels.tolist()) == 'FRJRSOSR'
```

<!-- grade_id: cell-d6f3d4de23bc3c91, points: 2, grade: true, solution: false -->
```python
import numpy as np
player_ids = np.array([p['player_id'] for p in team_roster])
unqiue_ids = np.unique(player_ids)
assert len(unqiue_ids) == len(team_roster)
```

### Question 9. Debug and Testing Functions

When you write a program, you have to make sure that it always works correctly even if some exteral data (e.g., the ESPN website) may be corrupted. For example, Some players in a certain team may not have their player no. Your code should be able tp catch such exceptions and replace the missing player no with the value "NA". (10 points)

<!-- grade_id: cell-87356ddb0aaeca8e, points: 0, grade: false, solution: true -->
```python
team_id = 258
team_name = "Virginia Cavaliers"
team_roster = extract_team_roster(team_id, team_name)
import numpy as np
positions = np.array([p['POS'] for p in team_roster])
position_label, position_counts = np.unique(positions, return_counts=True)
s = ''.join(position_label.tolist())
s

for dicts in team_roster:
    for key in dicts:
        if not dicts[key]:
            dicts[key] = "NA"
    print(dicts)
# raise NotImplementedError()
```

Output:
```
{'player_name': 'Ira Armstead', 'player_id': '4429219', 'player_no': '98', 'POS': 'QB', 'HT': '6\' 3"', 'WT': '200 lbs', 'Class': 'FR', 'Birthplace': 'South Bend, IN'}
{'player_name': 'Brennan Armstrong', 'player_id': '4370775', 'player_no': '5', 'POS': 'QB', 'HT': '6\' 2"', 'WT': '220 lbs', 'Class': 'SO', 'Birthplace': 'Shelby, OH'}
{'player_name': 'Jared Rayman', 'player_id': '4568497', 'player_no': 'NA', 'POS': 'QB', 'HT': '6\' 2"', 'WT': '180 lbs', 'Class': 'FR', 'Birthplace': 'Atlanta, GA'}
{'player_name': 'Lindell Stone', 'player_id': '4240850', 'player_no': '36', 'POS': 'QB', 'HT': '6\' 0"', 'WT': '230 lbs', 'Class': 'JR', 'Birthplace': 'Dallas, TX'}
{'player_name': 'Keytaon Thompson', 'player_id': '4242243', 'player_no': '99', 'POS': 'QB', 'HT': '6\' 4"', 'WT': '225 lbs', 'Class': 'JR', 'Birthplace': 'New Orleans, LA'}
{'player_name': 'Chris Barfield', 'player_id': '4684946', 'player_no': 'NA', 'POS': 'RB', 'HT': '5\' 10"', 'WT': '205 lbs', 'Class': 'FR', 'Birthplace': 'Richmon
```

<!-- grade_id: cell-0ed1c994a0f7210c, points: 5, grade: true, solution: false -->
```python
import numpy as np
positions = np.array([p['POS'] for p in team_roster])
position_label, position_counts = np.unique(positions, return_counts=True)
assert ''.join(position_label.tolist()) == 'CCBDBDEDLDTFBGLBLSNTOLOTPPKQBRBSTEWR'
```

<!-- grade_id: cell-cf43e6715aff64de, points: 5, grade: true, solution: false -->
```python
player_nos = np.array([p['player_no'] for p in team_roster])
na_count = len([no for no in player_nos if no == 'NA'])
assert na_count > 10
```

### Question 10. Add functions into a Python Modules

Now you have developed two solid functions: `get_team_url` and `extract_team_roster` and you would like to reuse them in other projects or share them with other people. Use the Notebook magic command to write the two functions with other neccessary  statements into a python file named `roster.py`. (4 points)

<!-- grade_id: cell-9121f8f48dd11851, points: 0, grade: false, solution: true -->
```python
%%writefile roster.py
import re
import requests
from bs4 import BeautifulSoup
def is_player_row(tag):
    if tag.name == 'tr' and tag.has_attr('data-idx'):
        return True
    return False

def get_team_url(team_id, team_name):
    """
    Get the team URL for a NCAA college football team
    
    Args:
        team_id: the id of a NACC football team assigned by ESPN
        team_name: the name of the football team
        
    Returns:
        team_url: a string representing the URL of the team
    """
    url_team_name = team_name.replace(" ","-").lower()
    team_url = 'https://www.espn.com/college-football/team/roster/_/id/%s/%s' % (team_id,url_team_name)
    return team_url

def extract_team_roster(team_id, team_name):
    """
    Extract the team roster for a NCAA college football team
    
    Args:
        team_id: the id of a NACC football team assigned by ESPN
        team_name: the name of the football team
        
    Returns:
        a list of dict representing the roster of the entire team.
    """
    roster = []
    url_team_name = team_name.replace(" ","-").lower()
    team_url = 'https://www.espn.com/college-football/team/roster/_/id/%s/%s' % (team_id,url_team_name)
    
    req = requests.get(team_url)
    soup = BeautifulSoup(req.content, 'html.parser')
    
    t = soup.find(class_="Table__Title").find_next("table")
    col_names = [column.get_text() for column in t.find_all('th')]

    player = {}
    players = []
    i = 0
    player_groups = soup.find_all(class_="Table__Title")
    
    for group_roster in player_groups:
        player_group_attr = soup.find(class_="Table__Title",string=group_roster.text).find_next('table')
        for c in player_group_attr.tbody.find_all(is_player_row):
            
            tds = c.find_all('td')
            urls = c.find_all('a', href=True)
            cols = []   
            for td in tds:
                cols.append(td.get_text())
            for url in urls:
                if (url['href']) not in cols:
                    cols.append(url['href'])
            players.append(cols)
            
            
            player_name = "".join(re.split("[^a-zA-Z ]*", players[i][1]))
            player_no = "".join(re.split("[^0-9]*", players[i][1]))

            player_url = players[i][7]
            player_id = "".join(re.split("[^0-9]*", player_url))
            player = {
                "player_name": player_name,
                "player_id": player_id,
                "player_no": player_no,
                "POS": players[i][2],
                "HT": players[i][3],
                "WT": players[i][4],
                "Class": players[i][5],
                "Birthplace": players[i][6]
            }
            roster.append(player)
            i = i + 1

    return roster
```

Output:
```
Writing roster.py
```

```python
!cat roster.py
```

Output:
```
import re
import requests
from bs4 import BeautifulSoup
def is_player_row(tag):
    if tag.name == 'tr' and tag.has_attr('data-idx'):
        return True
    return False

def get_team_url(team_id, team_name):
    """
    Get the team URL for a NCAA college football team
    
    Args:
        team_id: the id of a NACC football team assigned by ESPN
        team_name: the name of the football team
        
    Returns:
        team_url: a string representing the URL of the team
    """
    url_team_name = team_name.replace(" ","-").lower()
    team_url = 'https://www.espn.com/college-football/team/roster/_/id/%s/%s' % (team_id,url_team_name)
    return team_url

def extract_team_roster(team_id, team_name):
    """
    Extract the team roster for a NCAA college football team
    
    Args:
        team_id: the id of a NACC football team assigned by ESPN
        team_name: the name of the football team
        
    Returns:
        a list of dict represent
```

<!-- grade_id: cell-460dca90f00e20e6, points: 1, grade: true, solution: false -->
```python
import os
assert os.path.exists("roster.py")
```

<!-- grade_id: cell-d2d114880fe61a22, points: 3, grade: true, solution: false -->
```python
try:
    from bs4 import BeautifulSoup
    import requests
    from roster import get_team_url, extract_team_roster
    
    clemson_tiger_roster = extract_team_roster(258, "Virginia Cavaliers")
except:
    raise AssertionError("could not import")

```

### Congratulations! Yo have finished Part B of the problem.